# Working with multiple data sets

There are two data files that we'll be working with for this week's assignment.  They are described below.  Load those data files in with Pandas and then work to answering each of the questions below.

## npidata.csv

This file is basic information about every healthcare provider in the US.  It has one row for each NPI (National Provider Identifier).  It contains information such as the provider's name and address.


## cmsYYYY.csv

These are files about what kinds of procedures and patients providers in the US are serving under CMS programs, Medicare and Medicaid.  Each contains various statistics about providers over the course of a year.  There are three of these, for the years 2014, 2015, and 2016.  These files, however, don't contain information about the provider such as where the provider is located.

## Our Goals

For this assignment, we're going to want to compute some statistics based on the data in the **cms** files, but aggregate that data based on information in the **npidata** file.  As in last week's assignment, you'll need to store your answers in a variable called `answer` at the end of each step.

## Part 1

In this first step, we'll need to merge together all of the **cms** files into a single dataframe.  Be careful that these files might not be identical, so you'll have to look a little bit to figure out how to merge them.

As you are merging them, make sure that you retain information about which file (i.e. which year) the data came from.  Call that new columns `year`.

In your `answer` variable, provide a complete data frame that contains all of the rows and columns from the **cms** files, plus an additional column to store the year/file that particular row came from.

The assertion tests will give you a good idea as to if you're merging the files correctly.

<!-- grade_id: 01-solution, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import os
files = os.listdir('cms/')
```

<!-- grade_id: 01-solution, points: 0, grade: false, solution: true -->
```python
files
```

Output:
```
['cms_2016.csv', 'cms_2014.csv', 'cms_2015.csv', '.ipynb_checkpoints']
```

<!-- grade_id: 01-solution, points: 0, grade: false, solution: true -->
```python
dataframes = []
```

<!-- grade_id: 01-solution, points: 0, grade: false, solution: true -->
```python
for f in files:
    df = pd.read_csv(('cms/'+f),engine='python')
    dataframes.append(df) 
```

Output:
```
/opt/tljh/user/lib/python3.6/site-packages/pandas/io/parsers.py in __init__(self, f, **kwds)
   2291             encoding=self.encoding,
   2292             compression=self.compression,
-> 2293             memory_map=self.memory_map,
   2294         )
   2295         self.handles.extend(handles)

/opt/tljh/user/lib/python3.6/site-packages/pandas/io/common.py in _get_handle(path_or_buf, mode, encoding, compression, memory_map, is_text)
    400         elif is_text:
    401             # No explicit encoding
--> 402             f = open(path_or_buf, mode, errors="replace", newline="")
    403         else:
    404             # Binary mode

IsADirectoryError: [Errno 21] Is a directory: 'cms/.ipynb_checkpoints'
```

<!-- grade_id: 01-solution, points: 0, grade: false, solution: true -->
```python
len(dataframes)
```

<!-- grade_id: 01-solution, points: 0, grade: false, solution: true -->
```python
dataframes[1].head()
```

<!-- grade_id: 01-solution, points: 0, grade: false, solution: true -->
```python
len(dataframes) 
```

<!-- grade_id: 01-solution, points: 0, grade: false, solution: true -->
```python
cms = pd.concat(dataframes)
```

<!-- grade_id: 01-solution, points: 0, grade: false, solution: true -->
```python
len(cms)
```

<!-- grade_id: 01-solution, points: 0, grade: false, solution: true -->
```python
sum([len(x) for x in dataframes])
```

<!-- grade_id: 01-solution, points: 0, grade: false, solution: true -->
```python
cms2 = pd.concat(dataframes, keys=files)
```

<!-- grade_id: 01-solution, points: 0, grade: false, solution: true -->
```python
cms2.head().reset_index()
```

<!-- grade_id: 01-solution, points: 0, grade: false, solution: true -->
```python
cms2.index.levels[0]
```

<!-- grade_id: 01-solution, points: 0, grade: false, solution: true -->
```python
answer = cms2
```

<!-- grade_id: 01-tests, points: 5, grade: true, solution: false -->
```python
assert(answer.shape == (193862, 60))
assert(list(answer['year'].unique()) == [2014,2015,2016])
assert(set(answer.columns.str.lower()) == set(['year', 'nbr', 'npi', 'provider_type',
       'medicare_participation_indicator', 'number_of_hcpcs', 'total_services',
       'total_unique_benes', 'total_submitted_chrg_amt',
       'total_medicare_allowed_amt', 'total_medicare_payment_amt',
       'total_medicare_stnd_amt', 'drug_suppress_indicator',
       'number_of_drug_hcpcs', 'total_drug_services',
       'total_drug_unique_benes', 'total_drug_submitted_chrg_amt',
       'total_drug_medicare_allowed_amt', 'total_drug_medicare_payment_amt',
       'total_drug_medicare_stnd_amt', 'med_suppress_indicator',
       'number_of_med_hcpcs', 'total_med_services', 'total_med_unique_benes',
       'total_med_submitted_chrg_amt', 'total_med_medicare_allowed_amt',
       'total_med_medicare_payment_amt', 'total_med_medicare_stnd_amt',
       'beneficiary_average_age', 'beneficiary_age_less_65_count',
       'beneficiary_age_65_74_count', 'beneficiary_age_75_84_count',
       'beneficiary_age_greater_84_count', 'beneficiary_female_count',
       'beneficiary_male_count', 'beneficiary_race_white_count',
       'beneficiary_race_black_count', 'beneficiary_race_api_count',
       'beneficiary_race_hispanic_count', 'beneficiary_race_natind_count',
       'beneficiary_race_other_count', 'beneficiary_nondual_count',
       'beneficiary_dual_count', 'beneficiary_cc_afib_percent',
       'beneficiary_cc_alzrdsd_percent', 'beneficiary_cc_asthma_percent',
       'beneficiary_cc_cancer_percent', 'beneficiary_cc_chf_percent',
       'beneficiary_cc_ckd_percent', 'beneficiary_cc_copd_percent',
       'beneficiary_cc_depr_percent', 'beneficiary_cc_diab_percent',
       'beneficiary_cc_hyperl_percent', 'beneficiary_cc_hypert_percent',
       'beneficiary_cc_ihd_percent', 'beneficiary_cc_ost_percent',
       'beneficiary_cc_raoa_percent', 'beneficiary_cc_schiot_percent',
       'beneficiary_cc_strk_percent', 'beneficiary_average_risk_score']))
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-111-bdb8e920e30d> in <module>
----> 1 assert(answer.shape == (193862, 60))
      2 assert(list(answer['year'].unique()) == [2014,2015,2016])
      3 assert(set(answer.columns.str.lower()) == set(['year', 'nbr', 'npi', 'provider_type',
      4        'medicare_participation_indicator', 'number_of_hcpcs', 'total_services',
      5        'total_unique_benes', 'total_submitted_chrg_amt',

AssertionError:
```

## Part 2

In this next part, we're going to join the **cms** data with the provider information in the **/data/npidata.csv** file.  In this join, we don't want to lose any records from the **cms** files, even if no matching provider exists in the **npidata** file.  However, we don't care about any providers from the **npidata** file that don't have records in the **cms** files.  Those providers can be ignored.

Join the data files together to create one unified data frame called `answer`.  This dataframe should have all the columns from both **cms** and **npidata** files, joined together using the `npi` column.  

Note that `npi` is unique in the **npidata** file.

<!-- grade_id: 02-solution, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
```

<!-- grade_id: 02-solution, points: 0, grade: false, solution: true -->
```python
cms = cms2
npi = pd.read_csv('/data/npidata.csv')
```

<!-- grade_id: 02-solution, points: 0, grade: false, solution: true -->
```python
merged_df= npi.merge(cms, how='outer', left_on='NPI', right_on='NPI')
answer = merged_df
```

<!-- grade_id: 02-tests, points: 5, grade: true, solution: false -->
```python
import numpy
assert(answer.shape == (193862, 103))
assert(list(answer['Provider Business Mailing Address State Name'].unique()) == ['IL','MO',numpy.nan,'WY'])
assert(list(answer.groupby('Provider Business Mailing Address State Name').NPI.count()) == [111520, 53366, 4805])
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-124-6539db35ef6b> in <module>
      1 import numpy
----> 2 assert(answer.shape == (193862, 103))
      3 assert(list(answer['Provider Business Mailing Address State Name'].unique()) == ['IL','MO',numpy.nan,'WY'])
      4 assert(list(answer.groupby('Provider Business Mailing Address State Name').NPI.count()) == [111520, 53366, 4805])

AssertionError:
```

## Part 3

If you did everything right above, you'll notice that grouping by a column with NaN in it will cause some rows to disappear from the aggregation test.  So, let's create a new column called `'State'` that has the same value as whatever is in the `'Provider Business Mailing Address State Name'` column or a value of `'XX'` if there is no state information.

Set `answer` to be your final data frame with the new `'State'` column added.

<!-- grade_id: 03-solution, points: 0, grade: false, solution: true -->
```python
merged_df['State'] = State

group_by_state = merged_df.groupby('State')
answer = group_by_state

```

<!-- grade_id: 03-tests, points: 5, grade: true, solution: false -->
```python
assert(list(answer.groupby('State').npi.count()) == [111520, 53366, 4805, 24171])
assert(answer.shape == (193862, 104))
```

## Part 4

Next, let's summarize the data by year and by State.  Create a pivot table that contains one row for each state and one column for each year.  Within the pivot table, put a sum of total services as the values.

Assign `answer` to be that resulting pivot table.  In the tests, I'm going to plot a bar chart of your pivot table.

<!-- grade_id: 04-solution, points: 0, grade: false, solution: true -->
```python
%matplotlib inline
answer = group_by_state.groupby(['Year', 'State']).agg(
    {
        'Year':sum,
        'State': sum 
    }
```

<!-- grade_id: 04-tests, points: 5, grade: true, solution: false -->
```python
assert(answer.shape == (4,3))
assert(answer.sum().sum() == 519185664.6999999)
assert(answer[2016].sum() == 176596933.80000004)
assert(answer.loc['WY'].sum() == 10892707.4)
```

```python
%matplotlib inline
answer.plot.bar()
```