# Week 11 Exercises

_McKinney 7.2, 11.1_

The general activity that we are doing in this week's exercise is to explore and try to understand a particular data set.  In this case, it is 

1. Read in the data file and filter down to only looking at MO hospitals
2. Aggregate by Hospital: sum the Denominator, use min Start Date, use max End Date  (watch for rows with no Denominator value!)
3. Calculate the average per day across that entire span
4. Histogram
5. Rank and find the hospital with the most

## STEP 1 - Read and Filter

<img src="images/step1.png" alt="Read and Filter Output" style="width: 500px; float: right; margin-left: 20px; border: 1px solid">

In the first step, read in the data file from the data directory `/data/complications_all.csv`.  It is a CSV file and Pandas should read it in just fine.  Explore the file so that you understand the columns and values.  At the end of this step, create a variable called `mo_hospitals` that contains a data frame from the `complications_all.csv` file, filtered down to only contain those hospitals from the state of Missouri (MO).

A screenshot is included for reference.



<!-- grade_id: q1-solution, points: 0, grade: false, solution: true -->
```python
import pandas as pd
```

<!-- grade_id: q1-solution, points: 0, grade: false, solution: true -->
```python
data = pd.read_csv("complications_all.csv") 
data.head()
```

Output:
```
Facility ID                     Facility Name                 Address  \
0      010001  SOUTHEAST ALABAMA MEDICAL CENTER  1108 ROSS CLARK CIRCLE   
1      010001  SOUTHEAST ALABAMA MEDICAL CENTER  1108 ROSS CLARK CIRCLE   
2      010001  SOUTHEAST ALABAMA MEDICAL CENTER  1108 ROSS CLARK CIRCLE   
3      010001  SOUTHEAST ALABAMA MEDICAL CENTER  1108 ROSS CLARK CIRCLE   
4      010001  SOUTHEAST ALABAMA MEDICAL CENTER  1108 ROSS CLARK CIRCLE   

     City State  ZIP Code County Name    Phone Number     Measure ID  \
0  DOTHAN    AL     36301     HOUSTON  (334) 793-8701  COMP_HIP_KNEE   
1  DOTHAN    AL     36301     HOUSTON  (334) 793-8701    MORT_30_AMI   
2  DOTHAN    AL     36301     HOUSTON  (334) 793-8701   MORT_30_CABG   
3  DOTHAN    AL     36301     HOUSTON  (334) 793-8701   MORT_30_COPD   
4  DOTHAN    AL     36301     HOUSTON  (334) 793-8701     MORT_30_HF   

                                        Measure Name  \
0  Rate of complications for hip/knee replacement...   
1     
```

<!-- grade_id: q1-solution, points: 0, grade: false, solution: true -->
```python
filter = data['State']=='MO'
filter.head()
mo_hospitals = data[filter]
```

```python
mo_hospitals.head()
```

Output:
```
Facility ID          Facility Name        Address    City State  \
45534      260001  MERCY HOSPITAL JOPLIN  100 MERCY WAY  JOPLIN    MO   
45535      260001  MERCY HOSPITAL JOPLIN  100 MERCY WAY  JOPLIN    MO   
45536      260001  MERCY HOSPITAL JOPLIN  100 MERCY WAY  JOPLIN    MO   
45537      260001  MERCY HOSPITAL JOPLIN  100 MERCY WAY  JOPLIN    MO   
45538      260001  MERCY HOSPITAL JOPLIN  100 MERCY WAY  JOPLIN    MO   

       ZIP Code County Name    Phone Number     Measure ID  \
45534     64804      JASPER  (417) 781-2727  COMP_HIP_KNEE   
45535     64804      JASPER  (417) 781-2727    MORT_30_AMI   
45536     64804      JASPER  (417) 781-2727   MORT_30_CABG   
45537     64804      JASPER  (417) 781-2727   MORT_30_COPD   
45538     64804      JASPER  (417) 781-2727     MORT_30_HF   

                                            Measure Name  \
45534  Rate of complications for hip/knee replacement...   
45535               Death rate for heart attack patients   
45536         
```

<!-- grade_id: q1-tests, points: 2, grade: true, solution: false -->
```python
# These assertions will help make sure that you're on the right track.
assert(mo_hospitals['State'].unique() == ['MO'])
assert(mo_hospitals.shape == (2133,18))
```

## STEP 2 - Transform and Aggregate

<img src="images/step2.png" alt="Transforma and Aggregate Output" style="width: 500px; float: right; margin-left: 20px; border: 1px solid">

In the next step, we need to aggregate the results by hospital.  There are some key fields that we want to summarize, though:
* We want to know the earliest date that each hospital was participating in any program
* We want to know the latest date that each hospital stopped participating in any program
* We want to know the total number of patients in the denominators of these programs

Some things to note:
* You will need to convert the `Start Date` and `End Date` to actual datetime fields
* You will need to clean up and convert the `Denominator` field to just be numeric - the rule that you should use it to simply remove any records where the `Denominator` is `'Not Available'`

The final result of this step should be a new data frame called `mo_summary` that contains one row for each hospital and contains the min start date, max end date, and total denominator.  Use the names `start_date`, `end_date`, and `number` for those columns in `mo_summary`.

A screenshot is included for reference.

<!-- grade_id: q2-solution, points: 0, grade: false, solution: true -->
```python
import datetime

start_date = pd.to_datetime(data.int[0:8], format='%Y%m%d')
end_date = pd.to_datetime(data.int[0:8], format='%Y%m%d')
```

Output:
```
<ipython-input-87-5215485dc57e> in <module>
      1 import datetime
      2 
----> 3 start_date = pd.to_datetime(data.int[0:8], format='%Y%m%d')
      4 end_date = pd.to_datetime(data.int[0:8], format='%Y%m%d')

/opt/tljh/user/lib/python3.6/site-packages/pandas/core/generic.py in __getattr__(self, name)
   5272             if self._info_axis._can_hold_identifiers_and_holds_name(name):
   5273                 return self[name]
-> 5274             return object.__getattribute__(self, name)
   5275 
   5276     def __setattr__(self, name: str, value) -> None:

AttributeError: 'DataFrame' object has no attribute 'int'
```

<!-- grade_id: q2-solution, points: 0, grade: false, solution: true -->
```python
not_avail = mo_hospitals['Denominator'] == 'Not Available'
mo_hospitals[not_avail].shape
number = mo_hospitals[~not_avail]
number.shape
mo_summary = number.groupby(['Facility ID','Facility Name','Address']).min()
mo_summary.head()
```

Output:
```
City  \
Facility ID Facility Name                            Address                             
260001      MERCY HOSPITAL JOPLIN                    100 MERCY WAY              JOPLIN   
260004      COOPER COUNTY COMMUNITY HOSPITAL         17651 B HWY             BOONVILLE   
260005      SSM HEALTH ST JOSEPH HOSPITAL-ST CHARLES 300 1ST CAPITOL DR  SAINT CHARLES   
260006      MOSAIC LIFE CARE AT ST JOSEPH            5325 FARAON STREET   SAINT JOSEPH   
260009      BOTHWELL REGIONAL HEALTH CENTER          601 E 14TH ST             SEDALIA   

                                                                        State  \
Facility ID Facility Name                            Address                    
260001      MERCY HOSPITAL JOPLIN                    100 MERCY WAY         MO   
260004      COOPER COUNTY COMMUNITY HOSPITAL         17651 B HWY           MO   
260005      SSM HEALTH ST JOSEPH HOSPITAL-ST CHARLES 300 1ST CAPITOL DR    MO   
260006      MOSAIC LIFE CARE AT ST JOSEPH     
```

<!-- grade_id: q2-solution, points: 0, grade: false, solution: true -->
```python
mo_summary = pd.DataFrame()
```

<!-- grade_id: q2-test, points: 5, grade: true, solution: false -->
```python
assert(mo_summary['number'].sum() == 1766908)
assert(mo_summary['start_date'].min() == pd.Timestamp(2015,4,1))
assert(mo_summary['end_date'].max() == pd.Timestamp(2018,6,30))
assert(mo_summary.shape == (108,3))
assert(mo_summary.loc['BARNES JEWISH HOSPITAL'].number == 131313)
assert(mo_summary.loc['BOONE HOSPITAL CENTER'].number == 63099)
```

Output:
```
pandas/_libs/hashtable_class_helper.pxi in pandas._libs.hashtable.PyObjectHashTable.get_item()

pandas/_libs/hashtable_class_helper.pxi in pandas._libs.hashtable.PyObjectHashTable.get_item()

KeyError: 'number'
```

## STEP 3: Average Per Day

<img src="images/step3.png" alt="Average per Day" style="width: 500px; margin-left: 20px; float: right; border: 1px solid">


The next step, now that we have a start date, end date, and total patients for each day is to calculate how many patients on average per day this represents.  You will need to calculate the number of days between the start date and end date, and then the average as total patients divided by number of days.

Your final output should still be the variable `mo_summary`, and it will need to have two additional columns: `days` and `per_day`.

A screenshot is included for reference.


<!-- grade_id: q3-solution, points: 0, grade: false, solution: true -->
```python
# Put your code below and make sure that you add new columns to `mo_summary` 
# to have the contents described in the instructions.

mo_summary['duration'] = (mo_summary['End Date'] - mo_summary['Start Date'])
total_patient = mo_summary['Denominator'].sum

```

Output:
```
/opt/tljh/user/lib/python3.6/site-packages/pandas/core/ops/array_ops.py in na_arithmetic_op(left, right, op, str_rep)
    149         result = expressions.evaluate(op, str_rep, left, right)
    150     except TypeError:
--> 151         result = masked_arith_op(left, right, op)
    152 
    153     return missing.dispatch_fill_zeros(op, left, right, result)

/opt/tljh/user/lib/python3.6/site-packages/pandas/core/ops/array_ops.py in masked_arith_op(x, y, op)
     92         if mask.any():
     93             with np.errstate(all="ignore"):
---> 94                 result[mask] = op(xrav[mask], yrav[mask])
     95 
     96     else:

TypeError: unsupported operand type(s) for -: 'str' and 'str'
```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert(mo_summary['days'].sum() == 123718)
assert(mo_summary['per_day'].mean() == 13.841279932122697)
assert(mo_summary['per_day'].min() == 0.023744292237442923)
assert(mo_summary['days'].min() == 729)
assert(mo_summary['days'].max() == 1186)
```

## STEP 4: Histogram

<img src="images/step4.png" alt="Histogram" style="width: 300px; float: right; margin-left: 20px; border: 1px solid">

For this step, I want you to plot a histogram to give yourself a better understanding of this new metric we've calculated: how many patients per day have a potential complication.  Looking at the shape and distribution of the data might give us some indication of the hospital volume and quality.  (Don't take this too seriously, though.  We're cutting some important corners.)

A screenshot is included for reference.

<!-- grade_id: q4-solution, points: 3, grade: true, solution: true -->
```python
%matplotlib inline
# Use the built-in Pandas histogram plotting capability to plot a histogram of the `per_day` values.
import numpy as np
mo_summary['time'] = mo_summary.duration // np.timedelta64(1, 'Y')

mo_summary.time.hist()

```

Output:
```
<ipython-input-101-d7f293389d8c> in <module>
      2 # Use the built-in Pandas histogram plotting capability to plot a histogram of the `per_day` values.
      3 import numpy as np
----> 4 mo_summary['time'] = mo_summary.duration // np.timedelta64(1, 'Y')
      5 
      6 mo_summary.time.hist()

/opt/tljh/user/lib/python3.6/site-packages/pandas/core/generic.py in __getattr__(self, name)
   5272             if self._info_axis._can_hold_identifiers_and_holds_name(name):
   5273                 return self[name]
-> 5274             return object.__getattribute__(self, name)
   5275 
   5276     def __setattr__(self, name: str, value) -> None:

AttributeError: 'DataFrame' object has no attribute 'duration'
```

## STEP 5: Rank Order

Now that we see what the distribution looks like, you can tell that most of the `per_day` values are below 20.  There are, however, some outliers up around 100!  Who are those?  Let's rank the data set by the `per_day` value in descending order, examine them, and produce a list of the "top 3" based on this criteria.  That is, who are the top three hospitals based on having the highest `per_day` values.

At the end of this step, put those hospital names in a list called `top_hospitals` and that will be used for testing.  They should be in the order they appear `[#1, #2, #3]`.  That is, if the top hospitals were Mercy (99.3), BJC (97.2), and MoBap (90.1), then you would have `top_hospitals = ['Mercy','BJC','MoBap']`

<!-- grade_id: q5-solution, points: 0, grade: false, solution: true -->
```python
import numpy as np
```

<!-- grade_id: q5-solution, points: 0, grade: false, solution: true -->
```python
top_hospitals = ['BARNES JEWISH HOSPITAL', 'MERCY HOSPITAL SPRINGFIELD', 'COX MEDICAL CENTERS']

```

<!-- grade_id: q5-tests, points: 5, grade: true, solution: false -->
```python
assert(type(top_hospitals) == list)
### BEGIN HIDDEN TESTS
assert(top_hospitals == ['BARNES JEWISH HOSPITAL', 'MERCY HOSPITAL SPRINGFIELD', 'COX MEDICAL CENTERS'])
### END HIDDEN TESTS
```